# Aula 3 - Exercícios Extra - Python e Pandas
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
# Importe a biblioteca pandas e numpy
import pandas as pd
import numpy as np
```

### 3.8 Escreva uma função para ordenar uma lista de dicionários usando uma função Lambda. A função criada deve receber como argumento a chave do dicionário que será usada na ordenação da lista.


<!-- grade_id: cell-c705b9d11473b9f9, points: 0, grade: false, solution: true -->
```python
# Dica: use a função sorted. Use o argumento key da função sorted.
def sort_dict_list(l, chave):
    """retorne a lista ordenada
    """
    ### BEGIN SOLUTION
    return sorted(l, key = lambda x: x[chave])
    ### END SOLUTION

```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-9ffda88d58e9282a, points: 1, grade: true, solution: false -->
```python
# Validação
data = [{'nome': ['Pedro'], 'classe': ['V'], 'disciplina': ['portugues', 'matematica', 'ciencia']},
        {'nome': ['Jose'], 'classe': ['V'], 'disciplina': ['portugues', 'matematica', 'ciencia']},
       {'nome': ['Maria'], 'classe': ['V'], 'disciplina': ['portugues', 'matematica', 'ciencia']}]

data_sorted = [{'nome': ['Jose'], 'classe': ['V'], 'disciplina': ['portugues', 'matematica', 'ciencia']},
        {'nome': ['Maria'], 'classe': ['V'], 'disciplina': ['portugues', 'matematica', 'ciencia']},
       {'nome': ['Pedro'], 'classe': ['V'], 'disciplina': ['portugues', 'matematica', 'ciencia']}]

assert sort_dict_list(data, 'nome') == data_sorted
```

### 3.9 Escreva uma função para criar uma nova coluna chamada 'segundo' que tenha o segundo maior valor de cada linha de df.


<!-- grade_id: cell-553ee248eb14bc30, points: 0, grade: false, solution: true -->
```python
# Dica 1: Use a função apply do Dataframe aplicando uma função lambda.
#            e use a função sorted dentro da função lambda. Perceba que você precisa converter o argumento da
#            da função lambda para list antes de aplicar o método sorted.
# Dica 2: Aplique a função lambda no eixo 1.

def segundo_maior(df):
    """retorne o dataframe df com uma nova coluna 'segundo' que tenha o segundo maior valor de cada linha de df
    """
    ### BEGIN SOLUTION
    df["segundo"] = df.apply(lambda x: sorted(list(x))[-2], axis = 1)
    return df
    ### END SOLUTION

```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-0e2bt6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação
data = {'d1': [20, 36, 25,  4,  1, 20, 17, 47,  2, 40],
        'd2': [33, 25, 48, 39, 49, 37, 45, 36, 37, 31],
        'd3': [19, 47,  9, 27, 48, 26, 24, 21,  1,  2]}
df = pd.DataFrame(data)
assert list(segundo_maior(df)['segundo']) == [20, 36, 25, 27, 48, 26, 24, 36, 2, 31]
```