# Aula 1 - Exercícios - Sintaxe de Python e NumPy Array
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

### 1.1 Escreva uma função em Python que aceite um inteiro (n) de 1 a 9 e compute o valor de: n+nn+nnn
#### Valide os valores de entrada. Para isso, use a instrução raise ValueError
##### Exemplo: n=5. Então n + nn + nnn = 5 + 55 + 555 = 615

<!-- grade_id: cell-a2f686fdd79e4f51, points: 0, grade: false, solution: true -->
```python
def calculo_n(n):
    """retorne o valor de: n+nn+nnn
    """
    ### BEGIN SOLUTION
    
    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-5412543179b17057, points: 1, grade: true, solution: false -->
```python
# Validação
assert calculo_n(1) == 123
assert calculo_n(2) == 246
assert calculo_n(3) == 369
assert calculo_n(9) == 1107
```

### 1.2 Escreva uma função para retornar o maior item de uma lista de inteiros

<!-- grade_id: cell-2e044a5f0541da37, points: 0, grade: false, solution: true -->
```python
def maior(lista):
    """retorne o maior item da lista
    """
    ### BEGIN SOLUTION
    
    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-5bd4079b37b6a7f3, points: 1, grade: true, solution: false -->
```python
# Validação
assert maior([1,0,-1,100,7]) == 100
assert maior([-100,-1580,-1,-9100,-77]) == -1
```

### 1.3 Escreva uma função para contar a quantidade de strings, presentes numa lista, cujo tamanho é maior que 2 e o primeiro e último caracteres são os mesmos

#### Exemplo:  ['abc', 'xyz', 'aba', '1221']
#### Resultado esperado : 2


<!-- grade_id: cell-5420c3983940f212, points: 0, grade: false, solution: true -->
```python
# Dica 1: A função len retorna o tamanho da string. Exemplo: len('abc'), retorna 3
# Dica 2: Uma string é uma lista de caracteres
def contagem(lista):
    """retorne o maior item da lista
    """
    ### BEGIN SOLUTION

    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-f86ec02e2bcb84c6, points: 1, grade: true, solution: false -->
```python
# Validação
assert contagem(['abc', 'xyz', 'aba', '1221']) == 2
assert contagem(['ac', '1111', 'aasdfsdfdda', '12ewwww21']) == 3
assert contagem(['', 'xyzaaaaaaaaa', 'aa', 'babababa']) == 0
```

### 1.4 Crie um código python para imprimir o seguinte padrão. 
#### 1 
#### 1 2 
#### 1 2 3 
#### 1 2 3 4 
#### 1 2 3 4 5

```python
## Faça aqui

```

<!-- grade_id: cell-61d03fd9d1b0eb4e, points: 1, grade: true, solution: false -->
```python
### Este exercicio não tem assert

```

### 1.5 Escreva uma função Python para percorrer uma lista e retornar os números multiplos de N. Pare a iteração no ponto em que encontrar um número maior do que K. Não retorne o 1o número maior do que N.
#### Exemplo:  lista = [120, 115, 12, 32, 56, 200, 715, 152, 130, 155, 183, 300], N=5, K=200
#### Resultado Esperado: [120, 115, 200]

<!-- grade_id: cell-1ff4485b8100a753, points: 0, grade: false, solution: true -->
```python
def filtra_lista(lista, N, K):
    """retorne uma lista com todos os números multiplos de N menores do que K, 
       até o primeiro número maior do que K, não inclua este número na lista de resultado.
    """
    ### BEGIN SOLUTION
    
    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-b6c3a5db7e709b2a, points: 1, grade: true, solution: false -->
```python
# Validação
assert filtra_lista([120, 115, 12, 32, 56, 200, 715, 152, 130, 155, 183, 300], 5, 200) == [120, 115, 200]
```

### 1.6 Dado um número inteiro qualquer, escreva uma função python para retornar a quantidade de caracteres (dígitos) contida neste número.
##### Exemplo:  qtd_digitos(2345653) deve retornar 7

<!-- grade_id: cell-e69a127f71184645, points: 0, grade: false, solution: true -->
```python
def qtd_digitos(N):
    """retorne a quantidade de dígitos do número N.
    """
    ### BEGIN SOLUTION
    return
    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-418c099df61d368b, points: 1, grade: true, solution: false -->
```python
# Validação
assert qtd_digitos(2345653) == 7
```

### 1.7 Escreva uma função Python para somar os valores dos numeros presentes numa string de números separados por vírgula.
#### Exemplo:  "1.0,  4.0,  10.0,  8.0,  4.0,  20.0,  23.0"
#### Resultado esperado: 70

<!-- grade_id: cell-f6c50b3e165db993, points: 0, grade: false, solution: true -->
```python
def soma_string(texto):
    """retorne a soma dos numeros presentes na variavel texto
    """
    ### BEGIN SOLUTION
 
    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-401910203dadc4db, points: 1, grade: true, solution: false -->
```python
# Validação
var = "1.0,  4.0,  10.0,  8.0,  4.0,  20.0,  23.0"
soma_string(var) == 70.0
```

### 1.8 Crie um código python para imprimir o seguinte padrão. 
#### 5 4 3 2 1 
#### 4 3 2 1 
#### 3 2 1 
#### 2 1 
#### 1

```python
### Este exercicio não tem assert
## Faça aqui

```

### 1.9 Crie uma função python para inverter uma lista usando um loop do tipo for.
#### Exemplo:  inverter([120, 115, 12, 32, 56]) deve retornar [56, 32, 12, 115, 120]

<!-- grade_id: cell-513f2b547866e072, points: 0, grade: false, solution: true -->
```python
# Dica: A função ones do pacote numpy cria um array de 1s
def inverter(lista):
    """retorne uma versão invertida da lista passada como argumento
    """
    ### BEGIN SOLUTION

    ### END SOLUTION
```

```python
### Este exercicio não tem assert
## Faça aqui

```

### 1.10 Escreva uma função python para substituir vírgula por ponto numa string e depois converta os números para float. Retorne a lista de números.
#### Exemplo:  "0,20 0,40 0,56 0,99"
#### Resultado esperado: [0.2, 0.4, 0.56, 0.99]

<!-- grade_id: cell-e0a4e5deea9187c4, points: 0, grade: false, solution: true -->
```python
def convert_to_float(texto):
    """retorne uma lista de numeros do tipo float
    """
    ### BEGIN SOLUTION

    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-30d325468b6f6a67, points: 1, grade: true, solution: false -->
```python
# Validação
text = "0,20 0,40 0,56 0,99"
assert convert_to_float(text) == [0.2, 0.4, 0.56, 0.99]
```