# Aula 4 - Exercícios - Pandas, Fontes de dados e Agregação
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
## Fazer o download dos datasets
# Se já tiver feito o download, não precisa fazer novamente.
# Se aparecer um erro dizendo que a pasta não está vazia é porque o download já foi feito
!git clone https://github.com/alexlopespereira/curso_ciencia_dados2020.git ../../curso_ciencia_dados2020
```

```python
# Importe as bibliotecas pandas, numpy, glob e sqlalchemy
import pandas as pd
import numpy as np
import glob
import sqlalchemy as sqla
```

### 4.1 Carregue o arquivo do IDEB 2017 no formato de um DataFrame.  Crie uma função para calcular a média do IDEB de 2017 de todos os municípios de um determinado Estado. 
#### Remova os registros cujo valor na coluna REDE sejam Municipal, Estadual ou Federal. Deixe somente os registros cujos valores na coluna REDE sejam Pública.
#### Na célula de testes, crie um dataframe com o nome `df_ideb`.  Use os argumentos `skiprows`, `na_values` e  `skipfooter` para carregar a tabela corretamente.

<!-- grade_id: cell-c705ab9d11473b9f9, points: 0, grade: false, solution: true -->
```python
# Dica 1: Antes de calcular, certifique-se de interpretar os valores '-' como NA durante 
#         o carregamento dos dados ou os substitua por np.nan (NA) com o metodo replace.
# Dica 2: Use a função loc para fazer o filtro de registros.
# Dica 3: Dentre várias maneiras possíveis, uma maneira bastante concisa é usar o operador & (and) para
#         congregar dois critérios de filtro (seleção) de registros.

def media_ideb(df, sigla_estado):
    """retorne o valor da media do IDEB de 2017 contido na coluna col do Estado especificado por sigla_estado.
    """
    ### BEGIN SOLUTION

    ### END SOLUTION
```

```python
## Faça seus testes aqui

# path_ideb='https://github.com/alexlopespereira/curso_ciencia_dados2020/raw/master/data/originais/ideb/ideb_municipios2017.xlsx'
path_ideb='../../curso_ciencia_dados2020/data/originais/ideb/ideb_municipios2017.xlsx'

## Para que a validação na célula seguinte funcione, é preciso criar o dataframe chamado `df_ideb`
df_ideb = 

```

<!-- grade_id: cell-9ffdsa88d58e9282a, points: 1, grade: true, solution: false -->
```python
# Validação

###################################################
###       ATENÇÃO PARA ESTA INSTRUÇÃO           ###
# Crie na célula anterior o dataframe df_ideb.    #
# Use os argumentos skiprows, na_values e         #
# skipfooter para carregar a tabela corretamente. #
###################################################

assert round(media_ideb(df_ideb, 'DF'),1) == 3.4
assert round(media_ideb(df_ideb, 'SP'),1) == 4.0
assert round(media_ideb(df_ideb, 'GO'),1) == 4.2
```

### 4.2 Escreva uma função para salvar um DataFrame num banco de dados SQLite  

<!-- grade_id: cell-553qe5248eb14bc30, points: 0, grade: false, solution: true -->
```python
# Dica 1: A resolução deste exercício é bem semelhante ao código apresentado na Aula
# Dica 2: Antes de salvar no banco de dados, defina o nome do índice como index. Conforme exemplo abaixo:
#         df_idhm.index.names = ['index']

def save_to_sqlite(df, table_name, con_str):
    """salve o dataframe df num banco de dados sqlite especificado por db_path
    """
    ### BEGIN SOLUTION

    ### END SOLUTION
```

```python
## Faça seus testes aqui

# path_idhm = 'https://github.com/alexlopespereira/curso_ciencia_dados2020/raw/master/data/originais/idhm/idhm2013.xlsx'
path_idhm='../../curso_ciencia_dados2020/data/originais/idhm/idhm2013.xlsx'

```

<!-- grade_id: cell-0e2bh6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação
def load_table(table, con_str):
    con = sqla.create_engine(con_str)
    return pd.read_sql_query("SELECT * from {0}".format(table), con, index_col='index')
    
path_idhm='https://github.com/alexlopespereira/curso_ciencia_dados2020/raw/master/data/originais/idhm/idhm2013.xlsx'
con_str='sqlite:///idhm2013.sqlite'
table_name='IDHM'
df_idhm = pd.read_excel(path_idhm, sheet_name="IDHM")

df_idhm.index.names = ['index']
save_to_sqlite(df_idhm, table_name, con_str) # Função que deve ser implementada pelo aluno
idhm_from_db = load_table(table_name, con_str)

assert idhm_from_db.equals(df_idhm)
```

### 4.3 Escreva uma função executar uma query SQL de agregação num banco de dados SQLite que retorne a média dos valores de uma coluna `VALORES` agregados por uma coluna `CATEGORIAS`.
#### Dê o apelido `Media` para a métrica `AVG` na sua query. Ou seja: `AVG(CATEGORIA) AS Media`
#### Retorne a coluna `CATEGORIA` na tabela de resultados para identificar os valores. Ordene os resultados pelo valor da média em ordem decrescente.
#### Exemplo: Calcule média do IDHM (`VALORES='IDHM'`) agregado por Ano da pesquisa (`CATEGORIAS='ANO'`).

<!-- grade_id: cell-553e5248ebs14bc30, points: 0, grade: false, solution: true -->
```python
# Dica: A resposta deste exercício é bem semelhante ao código mostrado na aula, porém você
#       precisa parametrizar a string da query com os argumentos passados na função.
def soma_com_agregacao(val, cat, table, con_str):
    """retorne um dataframe contendo a soma dos valores em val dentro das categorias cat da tabela table
    """
    ### BEGIN SOLUTION

    ### END SOLUTION
```

```python
## Faça seus testes aqui

# path_idhm = 'https://github.com/alexlopespereira/curso_ciencia_dados2020/raw/master/data/originais/idhm/idhm2013.xlsx'
path_idhm='../../curso_ciencia_dados2020/data/originais/idhm/idhm2013.xlsx'

```

<!-- grade_id: cell-0e2bt6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação
path_idhm='https://github.com/alexlopespereira/curso_ciencia_dados2020/raw/master/data/originais/idhm/idhm2013.xlsx'
con_str='sqlite:///idhm2013.sqlite'
con = sqla.create_engine(con_str)
table_name='IDHM'
df_idhm = pd.read_excel(path_idhm, sheet_name="IDHM")
df_idhm.index.names = ['index']
df_idhm.to_sql(table_name, con, if_exists="replace")
test_data = {'ANO': {0: 2010, 1: 2000, 2: 1991}, 'Media': {0: 0.7, 1: 0.5, 2: 0.4}}
df_agg = soma_com_agregacao('IDHM', 'ANO', table_name, con_str)
df_agg['Media'] = df_agg['Media'].apply(lambda x: round(x, 1))
assert df_agg.to_dict() == test_data
```

### 4.4 Escreva uma função para fazer o merge dos dataframes de população e de PIB de 2017. A função deve receber como parâmetro o caminho para cada um dos dataframes.

<!-- grade_id: cell-553em5248eb14bc30, points: 0, grade: false, solution: true -->
```python
# Dica 1: A coluna que contem o código IBGE do município deve ter o mesmo formato (string) e a mesma
#         quantidade de caracteres nas duas tabelas.
# Dica 2: Converta o codigo do município para string ao carregar o dataframe de população utilizando
#         o mapeamento de tipos dtype={'cod_munic': str, 'cod_uf': str}. Assim o codigo do municipio não perde
#         os zeros a esquerda.
# Dica 3: Concatene o codigo da UF com o código do município usando o operador + e salve em outra coluna.
# Dica 4: Use o código apresentado na aula para carregar o dataframe do PIB de 2017, especifique o dtype da
#         coluna que contém o codigo do municipio.
# Dica 5: Renomeie as colunas apropriadamente antes de fazer o merge. A coluna com o nome do municipio deve se chamar
#         municipio (sem acento) e a coluna com o codigo ibge de 7 caracteres cod_ibge7. Veja o exemplo:
#         {'Unnamed: 0': 'nivel', 'Unnamed: 1': 'cod_ibge7', 'Unnamed: 2': 'municipio'}. O pandas transformará essas
#         colunas em municipio_x e municipio_y. Mantenha esses dois nomes de coluna, pois eles serão usados no exercício
#         3.7.
# Dica 6: Depois que os dois dataframes tiverem o codigo ibge (com valores equivalentes
#         e mesmo tipo de dados) você pode fazer o merge.

def merge_pop_pib(path_pop, path_pib):
    """retorne um dataframe resultante do merge dos dados do pib e da população
    """
    ### BEGIN SOLUTION

    ### END SOLUTION
```

```python
## Faça seus testes aqui
# path_pib = 'https://github.com/alexlopespereira/curso_ciencia_dados2020/raw/master/data/originais/pib/pib_municipios.xlsx'
# path_pop = 'https://github.com/alexlopespereira/curso_ciencia_dados2020/raw/master/data/originais/populacao/estimativa_dou_2017.xlsx'
path_pib = '../../curso_ciencia_dados2020/data/originais/pib/pib_municipios.xlsx'
path_pop='../../curso_ciencia_dados2020/data/originais/populacao/estimativa_dou_2017.xlsx'

```

<!-- grade_id: cell-0se2b6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação
path_pib = 'https://github.com/alexlopespereira/curso_ciencia_dados2020/raw/master/data/originais/pib/pib_municipios.xlsx'
path_pop = 'https://github.com/alexlopespereira/curso_ciencia_dados2020/raw/master/data/originais/populacao/estimativa_dou_2017.xlsx'

test_merge = {'pop2017': 25437, '2017': 498864}
df_pib_pop = merge_pop_pib(path_pop, path_pib)
assert df_pib_pop.loc[0,['pop2017', '2017']].to_dict() == test_merge 
```

### 4.5 Escreva uma função para calcular o PIB percapita de cada município do Brasil, considerando o dataframe do exercício anterior.
#### Nomeie a coluna do PIB Percapita com `pib_percapita`

<!-- grade_id: cell-553e524r8eb14bc30, points: 0, grade: false, solution: true -->
```python
def pib_percapita(pib_pop):
    """retorne retorne um dataframe contendo o codigo ibge de 7 caracteres,
       o nome do municipio e seu PIB percapita.
    """
    ### BEGIN SOLUTION

    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-0e2hb6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação
pib_percapita_2017 = pib_percapita(df_pib_pop)
assert round(pib_percapita(df_pib_pop).loc[0,'pib_percapita'],1) == 19.6
```

### 4.6 Escreva uma função para retornar os registros dos 10 maiores PIBs percapita do brasil (de 2017) a partir do dataframe do exercício anterior. 
#### Escreva uma função genérica que receba como parâmetro um dataframe e o nome da coluna de ordenação.


<!-- grade_id: cell-553e5s248eb14bc30, points: 0, grade: false, solution: true -->
```python
# Dica 1: Use o método sort_values para ordenar o dataframe. Para uma ordenação descendente use ascending=False.
# Dica 2: Use o atributo iloc para retornar todas as colunas dos 10 primeiros registros.

def top10(df, col):
    """retorne os top 10 registros do dataframe df ordenados por col
    """
    ### BEGIN SOLUTION

    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-0e2b6cefe06hbdc60, points: 1, grade: true, solution: false -->
```python
# Validação
test_data = {'cod_ibge7': {3677: '3536505', 3547: '3524709'},
             'municipio_x': {3677: 'Paulínia', 3547: 'Jaguariúna'},
             'uf': {3677: 'SP', 3547: 'SP'},
             'pib_percapita': {3677: 344.8, 3547: 209.3}}
test_data_y = {'cod_ibge7': {3677: '3536505', 3547: '3524709'},
             'municipio_y': {3677: 'Paulínia', 3547: 'Jaguariúna'},
             'uf': {3677: 'SP', 3547: 'SP'},
             'pib_percapita': {3677: 344.8, 3547: 209.3}}
cidades_0_9 = top10(pib_percapita_2017, 'pib_percapita')[['cod_ibge7', 'municipio_x', 'uf', 'pib_percapita']].iloc[[0,9]]
cidades_0_9_y = top10(pib_percapita_2017, 'pib_percapita')[['cod_ibge7', 'municipio_y', 'uf', 'pib_percapita']].iloc[[0,9]]

cidades_0_9['pib_percapita'] = cidades_0_9['pib_percapita'].apply(lambda x: round(x, 1))
cidades_0_9_y['pib_percapita'] = cidades_0_9_y['pib_percapita'].apply(lambda x: round(x, 1))

assert cidades_0_9.to_dict() == test_data or cidades_0_9_y.to_dict() == test_data_y
```

### 4.7 Escreva uma função para contar quantos nomes de municípios são diferentes nas bases de dados do PIB e da população. 
#### A coluna do nome dos municípios no dataset de PIB contém também o nome da UF. Remova o nome da UF, por exemplo, removendo os 5 últimos caracteres da string do nome do município, antes de efetuar a comparação.
#### Interprete este resultado. O que você pode extrair de aprendizado deste resultado? 

<!-- grade_id: cell-553e5248eb1t4bc30, points: 0, grade: false, solution: true -->
```python
# Dica 1: use a função map (element-wise) do objeto series, ou seja, sobre a coluna col2 do dataframe df.
# Dica 2: passe uma função lambda para o metodo map que executa as instruções do enunciado.
# Dica 3: compare o resultado da execução da funcao map com a coluna col1
# Dica 4: retorne a quantidade de diferenças resultantes da comparação
def compare_df_cols(df, col1, col2):
    """retorne a quantidade de registros em que o valor do registro na coluna col1 difere do valor em col2
    """
    ### BEGIN SOLUTION

    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-0e2b6cefet06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação
assert compare_df_cols(df_pib_pop, 'municipio_x', 'municipio_y') == 16
```

### 4.8 Considere algumas linhas de uma base de dados de morbidade apresentadas a seguir. Perceba que há dados de vários meses de 2020 para cada Estado da Federação.
#### Escreva uma função para pivotar esta base de dados apresentando os meses do ano nas colunas, os municípios nas linhas e a quantidade de óbitos como valores das células. 

|       | Município               |   cod_ibge6 | mes   |   ano | mes_ano    | Óbitos   |
|------:|:------------------------|------------:|:------|------:|:-----------|:---------|
|     0 | Alta Floresta D'Oeste   |      110001 | Abr   |  2020 | 2020-04-01 | 8        |
|     1 | Alto Alegre dos Parecis |      110037 | Abr   |  2020 | 2020-04-01 | 1        |
|     2 | Alto Paraíso            |      110040 | Abr   |  2020 | 2020-04-01 | 2        |
| 33347 | Vila Boa                |      522220 | Jan   |  2020 | 2020-01-01 | -        |
| 33348 | Vila Propício           |      522230 | Jan   |  2020 | 2020-01-01 | -        |
| 33349 | Brasília                |      530010 | Jan   |  2020 | 2020-01-01 | 578      |

#### O resultado esperado é semelhante a este:

|   cod_ibge6 | 2020-01-01   |   2020-02-01 |   2020-03-01 | 2020-04-01   | 2020-05-01   | 2020-06-01   |
|------------:|-------------:|-------------:|-------------:|-------------:|-------------:|-------------:|
|      110001 |  5 |            2 |            3 |  8 | 2            | 2            |
|      110002 | 12 |           13 |           16 | 10 | 8            | 7            |
|      110003 |  - |            1 |            1 |  - | -            | 2            |
|      110004 | 23 |            8 |           29 | 11 | 17           | 13           |
|      110005 |  1 |            2 |            2 |  3 | 1            | -            |

<!-- grade_id: cell-553e52r48eb14bc30, points: 0, grade: false, solution: true -->
```python
# Dica 1: Sempre dê preferência para fazer agregação usando códigos como identificador, 
#         use nomes em ultimo caso quando não houver códigos identificadores. 
# Dica 2: Se você usar o atributo `mes_ano` para pivotar, o pandas já ordena de forma cronológica as colunas.

def pivot_morbidade(file_path):
    """retorne um dataframe pivotado
    """
    ### BEGIN SOLUTION
    
    ### END SOLUTION
```

```python
## Faça seus testes aqui

# path_morbidade = 'https://github.com/alexlopespereira/curso_ciencia_dados2020/raw/master/data/originais/morbidade/consolidado/morbidade_simplificado.csv'
path_morbidade = '../../curso_ciencia_dados2020/data/originais/morbidade/consolidado/morbidade_simplificado.csv'

```

<!-- grade_id: cell-0e2b6tcefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação
path_morbidade = 'https://github.com/alexlopespereira/curso_ciencia_dados2020/raw/master/data/originais/morbidade/consolidado/morbidade_simplificado.csv'
test_data = {'2020-01-01': '5', '2020-02-01': '2', '2020-03-01': '3', '2020-04-01': '8', '2020-05-01': '2', 
             '2020-06-01': '2'}
assert pivot_morbidade(path_morbidade).loc[110001].to_dict() == test_data
```