# Aula 1 - Exercícios - Sintaxe de Python e NumPy Array
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

### 1.1 Escreva uma função em Python que aceite um inteiro (n) de 1 a 9 e compute o valor de: n+nn+nnn
#### Valide os valores de entrada. Para isso, use a instrução raise ValueError
##### Exemplo: n=5. Então n + nn + nnn = 5 + 55 + 555 = 615

<!-- grade_id: cell-a2f686fdd79e4f51, points: 0, grade: false, solution: true -->
```python
def calculo_n(n):
    """retorne o valor de: n+nn+nnn
    """
    ### BEGIN SOLUTION
    
    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-5412543179b17057, points: 1, grade: true, solution: false -->
```python
# Validação
assert calculo_n(1) == 123
assert calculo_n(2) == 246
assert calculo_n(3) == 369
assert calculo_n(9) == 1107
```

### 1.2 Escreva uma função para retornar o maior item de uma lista de inteiros

<!-- grade_id: cell-2e044a5f0541da37, points: 0, grade: false, solution: true -->
```python
def maior(lista):
    """retorne o maior item da lista
    """
    ### BEGIN SOLUTION
    
    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-5bd4079b37b6a7f3, points: 1, grade: true, solution: false -->
```python
# Validação
assert maior([1,0,-1,100,7]) == 100
assert maior([-100,-1580,-1,-9100,-77]) == -1
```

### 1.3 Escreva uma função para contar a quantidade de strings, presentes numa lista, cujo tamanho é maior que 2 e o primeiro e último caracteres são os mesmos

#### Exemplo:  ['abc', 'xyz', 'aba', '1221']
#### Resultado esperado : 2


<!-- grade_id: cell-5420c3983940f212, points: 0, grade: false, solution: true -->
```python
# Dica 1: A função len retorna o tamanho da string. Exemplo: len('abc'), retorna 3
# Dica 2: Uma string é uma lista de caracteres
def contagem(lista):
    """retorne o maior item da lista
    """
    ### BEGIN SOLUTION

    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-f86ec02e2bcb84c6, points: 1, grade: true, solution: false -->
```python
# Validação
assert contagem(['abc', 'xyz', 'aba', '1221']) == 2
assert contagem(['ac', '1111', 'aasdfsdfdda', '12ewwww21']) == 3
assert contagem(['', 'xyzaaaaaaaaa', 'aa', 'babababa']) == 0
```

### 1.4 Escreva uma função Python para gerar um dicionário que contenha pares (chave e valor) de números. A chave deve variar de 1 a n, e os valores sejam o valor da chave elevado ao quadrado. 
#### Exemplo:  Para n = 5
#### Resultado Esperado: {1: 1, 2: 4, 3: 9, 4: 16, 5: 25}

<!-- grade_id: cell-bb9537795a06bd7e, points: 0, grade: false, solution: true -->
```python
# Dica: use a função range para criar a sequencia de numeros desejados
def create_dict(n):
    """retorne um dicionário com pares (x, x*x)
    """
    ### BEGIN SOLUTION
    
    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-61d03fd9d1b0eb4e, points: 1, grade: true, solution: false -->
```python
# Validação
assert create_dict(5) == {1: 1, 2: 4, 3: 9, 4: 16, 5: 25}
```

### 1.5 Escreva uma função Python para remover itens duplicados de um dicionário, ou seja, cujos valores estejam duplicados.
#### Exemplo:  d =  {1: 1, 2: 1, 3: 9, 4: 1, 5: 25}
#### Resultado Esperado: {1: 1, 3: 9, 5: 25}

<!-- grade_id: cell-1ff4485b8100a753, points: 0, grade: false, solution: true -->
```python
def remove_duplicates(data):
    """retorne um dicionário sem duplicatas
    """
    ### BEGIN SOLUTION
    
    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-b6c3a5db7e709b2a, points: 1, grade: true, solution: false -->
```python
# Validação
assert remove_duplicates({1: 1, 2: 1, 3: 9, 4: 1, 5: 25}) ==  {1: 1, 3: 9, 5: 25}
data = {'id1': {'nome': ['Sara'], 'classe': ['A'], 'disciplina': ['portugues, historia, ciencia']},
        'id2': {'nome': ['Joao'], 'classe': ['A'], 'disciplina': ['portugues, historia, ciencia']},
        'id3': {'nome': ['Sara'], 'classe': ['A'], 'disciplina': ['portugues, historia, ciencia']},
        'id4': {'nome': ['Maria'], 'classe': ['A'], 'disciplina': ['portugues, historia, ciencia']}
}
no_duplicates = {'id1': {'nome': ['Sara'], 'classe': ['A'], 'disciplina': ['portugues, historia, ciencia']}, 
                 'id2': {'nome': ['Joao'], 'classe': ['A'], 'disciplina': ['portugues, historia, ciencia']}, 
                 'id4': {'nome': ['Maria'], 'classe': ['A'], 'disciplina': ['portugues, historia, ciencia']}}
assert remove_duplicates(data) == no_duplicates
```

### 1.6 Escreva uma função Python para combinar os itens de mesma chave de uma lista de  dicionários, somando os valores correspondentes.
##### Exemplo:  para [{'item': 'carro', 'quantidade': 400}, {'item': 'moto', 'quantidade': 300}, {'item': 'carro', 'quantidade': 750}]
##### Resultado esperado: {'carro': 1150, 'moto': 300}

<!-- grade_id: cell-e69a127f71184645, points: 0, grade: false, solution: true -->
```python
def combine_items(item_list):
    """retorne um dicionário combinando (somando) as quantidades do mesmo item
    """
    ### BEGIN SOLUTION

    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-418c099df61d368b, points: 1, grade: true, solution: false -->
```python
# Validação
input_data = [{'item': 'carro', 'quantidade': 400}, 
              {'item': 'moto', 'quantidade': 300}, 
              {'item': 'carro', 'quantidade': 750}]
output_result = {'carro': 1150, 'moto': 300}
assert combine_items(input_data) == output_result
```

### 1.7 Escreva uma função Python para somar os valores dos numeros presentes numa string de números separados por vírgula.
#### Exemplo:  "1.0,  4.0,  10.0,  8.0,  4.0,  20.0,  23.0"
#### Resultado esperado: 70

<!-- grade_id: cell-f6c50b3e165db993, points: 0, grade: false, solution: true -->
```python
def soma_string(texto):
    """retorne a soma dos numeros presentes na variavel texto
    """
    ### BEGIN SOLUTION
 
    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-401910203dadc4db, points: 1, grade: true, solution: false -->
```python
# Validação
var = "1.0,  4.0,  10.0,  8.0,  4.0,  20.0,  23.0"
soma_string(var) == 70.0
```

### 1.8 Usando um defaultdict, escreva uma função Python para transformar uma lista de listas num dicionário. Os valores repetidos de uma lista devem ser associados a mesma chave do dicionário, como no exemplo abaixo.
#### Exemplo:  [ ["carro", 10], ["moto", 20], ["trem", 5], ["carro", 30], ["caminhao", 20] ]
#### Resultado esperado: {"carro": [10, 20], "moto": [20], "trem": [5], "caminhao": [20] }

<!-- grade_id: cell-f7b05dc12adbf2e2, points: 0, grade: false, solution: true -->
```python
def dict_from_list(lista):
    """retorne o valor maximo dentre os resultados da soma por linha de um array 2D 
    """
    ### BEGIN SOLUTION

    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-af1b0435b6aacc9b, points: 1, grade: true, solution: false -->
```python
# Validação
lista2 = [ ["carro", 10], ["moto", 20], ["trem", 5], ["carro", 30], ["caminhao", 20] ]
assert dict_from_list(lista2) == {'carro': [10, 30], 'moto': [20], 'trem': [5], 'caminhao': [20]}
```

### 1.9 Escreva uma função Python para computar o valor total de uma compra de supermercado e dar baixa no estoque (deduzir a quantidade comprada da quantidade em estoque). Retorne o valor total da compra e altere o dicionário de estoque sem retorná-lo.
#### Se a quantidade do item em estoque for menor do que a quantidade da lista de compra, o estoque restante deve ser vendido.

#### Exemplo:  lista_compras = {"arroz":1, "feijao":2, "macarrao":3}
####                   estoque = {"arroz": 60, "feijao": 30, "macarrao": 0, "farinha": 15}
####                   precos = {"arroz": 4, "feijao": 2, "macarrao": 1.5, "farinha": 3}
#### Resultado esperado: total = 6

<!-- grade_id: cell-513f2b547866e072, points: 0, grade: false, solution: true -->
```python
# Dica: A função ones do pacote numpy cria um array de 1s
def compra(estoque, precos, lista):
    """retorne um array 2d preenchido com 1s nas bordas e 0s no interior 
    """
    ### BEGIN SOLUTION

    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-73baa7b9cc0314dc, points: 1, grade: true, solution: false -->
```python
# Validação
lista_compras = {"arroz":1, "feijao":2, "macarrao":3}
estoque = {"arroz": 60, "feijao": 30, "macarrao": 1, "farinha": 15}
precos = {"arroz": 4, "feijao": 2, "macarrao": 1.5, "farinha": 3}
compra(estoque, precos, lista_compras) == 9.5
```

### 1.10 Escreva uma função python para substituir vírgula por ponto numa string e depois converta os números para float. Retorne a lista de números.
#### Exemplo:  "0,20 0,40 0,56 0,99"
#### Resultado esperado: [0.2, 0.4, 0.56, 0.99]

<!-- grade_id: cell-e0a4e5deea9187c4, points: 0, grade: false, solution: true -->
```python
# Dica: T(°C) = (T(°F) - 32) × 5/9
def convert_to_float(texto):
    """retorne uma lista de numeros do tipo float
    """
    ### BEGIN SOLUTION

    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-30d325468b6f6a67, points: 1, grade: true, solution: false -->
```python
# Validação
text = "0,20 0,40 0,56 0,99"
assert convert_to_float(text) == [0.2, 0.4, 0.56, 0.99]
```