# Aula 2 - Exercícios Extras - Sintaxe de Python e NumPy Array

```python
import numpy as np
```

### 2.11 Escreva uma função Python para selecionar dentre os elementos de um array aquele que mais se aproxima de um determinado valor N passado como argumento da função.
#### Se houver dois elementos equidistantes de N e forem os mais próximos de N, retorne o menor. Ou seja, para [1, 3, 4] e N=2, retorne 1.
#### Não use laços (loops), use funções Numpy.
#### Exemplo:  para um array([0, 12, 45, 34, 99]) e N=20
#### Resultado esperado: 12

<!-- grade_id: cell-e0a4e5deea9187c4, points: 0, grade: false, solution: true -->
```python
# Dica: Use as funções np.abs, ndarray.argmin e ndarray.sort 
def get_closest(arr, N):
    """retorne o elemento do array que mais se aproxima de N.
       No caso de dois elementos equidistantes, retorne o menor.
    """
    ### BEGIN SOLUTION

    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-30d325468b6f6a67, points: 1, grade: true, solution: false -->
```python
# Validação
inarr = np.array([3, 6, 1, 4])
assert get_closest(inarr, 2) == 1
```

### 2.12 Escreva uma função que receba uma lista de inteiros como parâmetro e retorne um dicionário com as chaves sendo os elementos da lista, e os valores sendo os elementos da mesma lista, mas na ordem inversa.
#### Exemplo: para a lista [1, 2, 3, 4, 5]
#### Resultado esperado: {1: 5, 2: 4, 3: 3, 4: 2, 5: 1}

```python
# Dica: use a função zip. Ela retorna itens agrupados dois a dois.
def create_dict(lista):
    """retorne o maior item da lista
    """
    ### BEGIN SOLUTION

    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

```python
# Validação
assert create_dict([1, 2, 3, 4, 5]) == {1: 5, 2: 4, 3: 3, 4: 2, 5: 1}
assert create_dict([-100,-1580,-1,-9100,-77]) == {-100: -77, -1580: -9100, -1: -1, -9100: -1580, -77: -100}
```