# Aula 2 - Exercícios - Sintaxe de Python e NumPy Array
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
# import numpy
import numpy as np
```

### 2.1 Escreva uma função Python para contar a quantidade de itens de um dado valor num array NumPy
#### Exemplo:  para um array([1  4  10  8  4  20  23]) e n=4
#### Resultado esperado: 2
#### Obs.: Lembre-se de usar funções numpy que realizam as operações sem a necessidade de escrever laços (loops)

<!-- grade_id: cell-f6c50b3e165db993, points: 0, grade: false, solution: true -->
```python
# Dica: A função count_nonzero do pacote numpy conta a quantidade 
#       de elementos diferentes de zero ou a quantidade de elementos True.

def count_value(array, n):
    """retorne a quantidade de vezes que n ocorre 
    """
    ### BEGIN SOLUTION
    return np.count_nonzero(array == n)
    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-401910203dadc4db, points: 1, grade: true, solution: false -->
```python
# Validação
a1 = np.array([1,0,1,9,5,3,1,2,4,5,6,1,2,3,4,8,9,2,1])
assert count_value(a1, 1) == 5

```

### 2.2 Escreva uma função Python para somar os valores de cada linha de um array numpy 2D e depois retornar a maior destas somas
#### Exemplo:  para um array([[1  4 ], [10  8], [4  20]]) 
#### Resultado esperado: 24

<!-- grade_id: cell-f7b05dc12adbf2e2, points: 0, grade: false, solution: true -->
```python
# Dica: A função sum do numpy pode fazer o calculo da soma por linha ou por coluna, 
#       desde que o parâmetro axis seja especificado corretamente.
def max_sum(x):
    """retorne o valor maximo dentre os resultados da soma por linha de um array 2D 
    """
    ### BEGIN SOLUTION
    total_by_line = np.sum(x, axis=1)
    return np.max(total_by_line)
    ### END SOLUTION
```

```python
## Faça seus testes aqui


```

<!-- grade_id: cell-af1b0435b6aacc9b, points: 1, grade: true, solution: false -->
```python
# Validação
a1 = np.array([[1, 4 ], [10, 8], [4, 20]])
a2 = np.array([[1,0,1,9,5],[3,1,2,4,5],[6,1,2,3,4],[8,9,2,1,7]])
assert max_sum(a1) == 24
assert max_sum(a2) == 27
```

### 2.3 Escreva uma função Python para criar um array numpy 2D com 1s nos elementos das bordas e com 0s no interior. O tamanho do array deve ser parametrizável.
#### Exemplo:  para um n=3
#### Resultado esperado: array([[1, 1, 1], [1, 0, 1], [1, 1, 1]])

<!-- grade_id: cell-513f2b547866e072, points: 0, grade: false, solution: true -->
```python
# Dica: A função ones do pacote numpy cria um array de 1s
def create_array(n):
    """retorne um array 2d preenchido com 1s nas bordas e 0s no interior 
    """
    ### BEGIN SOLUTION
    x = np.ones((n,n))
    x[1:-1,1:-1] = 0
    return x
    ### END SOLUTION
```

```python
## Faça seus testes aqui


```

<!-- grade_id: cell-73baa7b9cc0314dc, points: 1, grade: true, solution: false -->
```python
# Validação
a3 = np.array([[1, 1, 1], [1, 0, 1], [1, 1, 1]])
a6 = np.array([[1., 1., 1., 1., 1., 1.],
       [1., 0., 0., 0., 0., 1.],
       [1., 0., 0., 0., 0., 1.],
       [1., 0., 0., 0., 0., 1.],
       [1., 0., 0., 0., 0., 1.],
       [1., 1., 1., 1., 1., 1.]])
assert (create_array(3) == a3).all()
assert (create_array(6) == a6).all()
```

### 2.4 Escreva uma função Python para converter para celcius um array de temperaturas em Fahrenheit. Arredonde o resultado até 2a casa decimal usando a função .round(decimals=2) do numpy
#### Exemplo:  para um array([0, 12, 45.21, 34, 99.91])
#### Resultado esperado: array([-17.77, -11.11, 7.33, 1.11, 37.72])

<!-- grade_id: cell-e0a4e5deea9187c4, points: 0, grade: false, solution: true -->
```python
# Dica: T(°C) = (T(°F) - 32) × 5/9
def get_ceusius(F):
    """retorne um array contendo as temperaturas em graus Ceusius
    """
    ### BEGIN SOLUTION
    return ((F-32)*5/9).round(decimals=2)
    ### END SOLUTION
```

```python
## Faça seus testes aqui


```

<!-- grade_id: cell-30d325468b6f6a67, points: 1, grade: true, solution: false -->
```python
# Validação
tf = np.array([0, 12, 45.21, 34, 99.91])
tc = np.array([-17.78, -11.11, 7.34, 1.11, 37.73])
assert (get_ceusius(tf) == tc).all()
```

### 2.5 Escreva uma função Python para selecionar os elementos de um array que são ao mesmo tempo: maiores que 10, menores que 50 e multiplos de 3.
#### Exemplo:  para um array([0, 12, 45, 34, 99])
#### Resultado esperado: array([12, 45])

<!-- grade_id: cell-e0a4e5deea9187c4, points: 0, grade: false, solution: true -->
```python
# Dica: Combine a explicação do slide 29 da aula2 com a sua solução
#       do exercício W2.22 
def filter_array(arr):
    """retorne um array numpy contendo números maiores que 10, 
       menores que 50 e multiplos de 3
    """
    ### BEGIN SOLUTION
    return arr[(arr > 10) & (arr < 50) & (arr%3==0)]
    ### END SOLUTION
```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-30d325468b6f6a67, points: 1, grade: true, solution: false -->
```python
# Validação
inarr = np.array([0, 12, 45, 34, 99])
outarr = np.array([12, 45])
assert (filter_array(inarr) == outarr).all()
```